***Perceptron***


```python
import pandas as pd
import numpy as np
from numpy import array, exp
import matplotlib.pyplot as plt
%matplotlib inline
from sklearn.datasets import make_gaussian_quantiles
from sklearn.neighbors import KNeighborsClassifier
```

Use the `make_guassian_quantiles` function from the `sklearn.datasets`class to create 1000 2-dimensional data points with two labels. Put the data in the variable `X` and the classes in the variable `c`. Use `random_state=42`.

<!-- grade_id: cell-63d28a4a996eb894, points: 0, grade: false, solution: true -->
```python
#X = make_gaussian_quantiles(random_state=42)
df= make_gaussian_quantiles(random_state=42, n_samples=1000, n_classes=2)
X = df[0]
c = df[1]




```

<!-- grade_id: cell-966bf7e3648a388c, points: 1, grade: true, solution: false -->
```python
assert X.shape == (1000,2)
assert c.shape == (1000,)
assert c.max() == 1
assert c.min() == 0
```

Make a scatterplot of the generated points using two different colors for the different classes.

<!-- grade_id: cell-7cad454d15726ed1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
color = ['r', 'b']




```

Create a function that takes as input `x` and returns the sigmoid of `x` and test it on the given vector `v` (storing the result in the variable `v_activated`.

*N.B.* The sigmoid function should accept any iterable as input, but must return a numpy array.

<!-- grade_id: cell-25c6e2ab20993a7e, points: 0, grade: false, solution: true -->
```python
v = [1, 2, -1]

def sigmoid(x):
    y = np.array(x)
    return 1 / (1 + exp(-y))
    
    


# YOUR CODE HERE
v_activated = sigmoid(v)
```

<!-- grade_id: cell-083dec44f74fed1a, points: 1, grade: true, solution: false -->
```python
assert np.round(v_activated[0], 2) == 0.73
assert np.round(v_activated[1], 2) == 0.88
assert np.round(v_activated[2], 2) == 0.27


```

Define the step function. As for the sigmoid, create a function that takes as input `x` and returns the sigmoid of `x` and test it on the given vector `v` (storing the result in the variable `v_activated`.

<!-- grade_id: cell-4f9ebef5da14c733, points: 0, grade: false, solution: true -->
```python
v = [1, 0, -1]

def step(x):
    y = np.array(x)
    return 1/(1+exp(-y))

v_activated = step(v)

# YOUR CODE HERE
v_activated = np.round(v_activated, 0)
```

```python
v_activated
```

Output:
```
array([1., 0., 0.])
```

<!-- grade_id: cell-978a2c6edd8b6234, points: 0, grade: true, solution: false -->
```python
assert v_activated[0] == 1.0
assert v_activated[1] == 1.0
assert v_activated [2] == 0.0
```

Output:
```
AssertionError                            Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_6500/2198280594.py in <module>
      1 assert v_activated[0] == 1.0
----> 2 assert v_activated[1] == 1.0
      3 assert v_activated [2] == 0.0

AssertionError:
```

You want to build a model to choose which book to read based on the following three variables:

- average reviews on Goodreads (from 1 to 5)
- number of 1 stars reviews on Goodreads
- number of reviews on Goodreads
- it's fiction or not

Create a potential dataset with reasonable examples with 5 data points presented in a numpy array of the form `X = np.array([[f11,f12...], [f21, f22,...],...,[f51,f52,...]])` so that the shape of X will be (3,4). In addition, create the `y` vector with binary labels.

<!-- grade_id: cell-e42eb35e7fbdece3, points: 0, grade: false, solution: true -->
```python
X = np.array([[2, 4 , 5, 1], [2, 3, 5, 1], [1, 4, 5, 0]])
y = np.random.randint(2, size = 3)
X.shape
```

Output:
```
(3, 4)
```

<!-- grade_id: cell-999c065e8557524e, points: 0, grade: true, solution: false -->
```python
assert X.shape == (3,4)
assert np.array([1<=x[0]<=5 for x in X]).sum() == 3
assert np.array([x[3] for x in X]).max() == 1
assert np.array([x[3] for x in X]).min() == 0
assert np.array([x[1]<=x[2] for x in X]).sum() == 3
assert y.max() == 1
assert y.min() == 0
assert y.shape == (3,)


```

You want to define a boundary n-1 hyperplane for the data you have generated. Create the vector W containing the weights of your model initialized by following a normal distribution (you can use `numpy.random.normal`. Don't forget the bias `b`.

<!-- grade_id: cell-3ada8bc49ebdd9aa, points: 0, grade: false, solution: true -->
```python
W = 0
b = 0

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c04231e0b720cac4, points: 1, grade: true, solution: false -->
```python
assert W.shape == (4,) or W.shape == (4, 1)
assert b.shape == (1,) or b.shape == (1, 1)
```

Create an array containing the score of your random model on each of the data points you generated. Store the result in the `scores` variable. 

*Hint*: Use the dot product. I recommend you to reshape it at the end so that it has only one dimension, not two.

<!-- grade_id: cell-951dd174a94f79aa, points: 0, grade: false, solution: true -->
```python
scores = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8ad99d3f33895cd4, points: 0, grade: true, solution: false -->
```python
assert scores.shape == (3,)

```

Activate all the scores by using the step function you defined before and store them in the `step_scores` variable.

<!-- grade_id: cell-3f689055ccafa454, points: 0, grade: false, solution: true -->
```python
step_scores = 0

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1f0db6a939917418, points: 0, grade: true, solution: false -->
```python
assert step_scores.shape == (3,)
```

Activate all the scores by using the sigmoid function you defined before and store them in the `sigmoid_scores` variable.

<!-- grade_id: cell-f765e664087d74f0, points: 0, grade: false, solution: true -->
```python
sigmoid_scores = 0

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4210905859c1bf9d, points: 0, grade: true, solution: false -->
```python
assert sigmoid_scores.shape == (3,)
```

Take the `step_scores` as predicted scores. You know that you can update your model weights using the formula:
$$w_{i(new)}=w_{i(old)} + \eta(y-\hat{y})x_{i}$$
where $y$ is the correct label of your data point (your `y[i]`) and $\hat{y}$ is the `step_scores`. And what about $\eta$? It's the learning rate. 

Set the learning rate to 0.1, and create the function `update(W_old, b_old, X, y, step_scores, eta)` that returns the updated weigths vector and bias. Test it using as input (X,y) the first data point of your dataset (*Hint*:`X[0], y[0]`, and of course `step_score[0]`).

<!-- grade_id: cell-d26387e4b70d46a5, points: 0, grade: false, solution: true -->
```python
def update_weights(W_old, b_old, X, y, step_scores, eta):
    return 0, 0

W_new, b_new = update_weights(0,0,0,0,0,0)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8465db2dfb20da02, points: 0, grade: true, solution: false -->
```python
assert W_new.shape == (4,)
assert b_new.shape == (1,)
```

Create a for loop to update the weights on each of the points in your dataset. In other words:
Iterate over the data points in X and the labels in y
- compute the `step_score` using the W, b that you have randomly inititated
- use the `update_weights` function to update them
- repeat

Do it for 10 epochs.

<!-- grade_id: cell-67c5d7eaf5a420b2, points: 0, grade: false, solution: true -->
```python
epochs = 0
step_score = 0
eta = 0
for epoch in range(epochs):
    for x, label in zip(X, y):
        pass

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2c08018e859be627, points: 1, grade: true, solution: false -->
```python
assert W_new.shape == (4,)
assert b_new.shape == (1,)
```

Now that you got so far, let's try to train a model for the Iris Dataset.

Maybe you don't like it, but you can't do great things using a perceptron. At the same time, understanding the steps behind it can help you a lot to build your understanding around the neural networks. 

In fact, a neural network is basically the combination of multiple perceptrons on the same input, and a deep neural network is a combination of multiple perceptrons on the combination of multiple perceptrons on ... on the same input! 

I'll help you giving an helper function that, given weights, bias, (x1,x2) coordinates and binary labels, plots the data points and the boundary line:

```python
from IPython import display
import time
# function that plots data and linear classifier
def plot_border(w1, w2, b, data, fig):
    
    plt.axis([0, 10, 0, 6])
    plt.grid()
    
    # scatter data
    for i in range(len(data)) :
        point = data[i]
        color = "r"
        if point[2] == 0 :
            color = "b"
        plt.scatter(point[0], point[1], c=color)

    # separation line
    x = np.linspace(0, 10, 100)
    plt.plot(x, -w1/w2*x+b/w2, '-g', label='y=w1*x+b')
    plt.show()
```

```python
### load data ###

# each point is a tuple (sepal length, sepal width, flower type)
# types: 0 for Setosa & 1 for Versicolor

data = [[ 5.1,  3.5,  0 ],
       [ 4.9,  3. ,  0 ],
       [ 4.7,  3.2,  0 ],
       [ 4.6,  3.1,  0 ],
       [ 5. ,  3.6,  0 ],
       [ 5.4,  3.9,  0 ],
       [ 4.6,  3.4,  0 ],
       [ 5. ,  3.4,  0 ],
       [ 4.4,  2.9,  0 ],
       [ 4.9,  3.1,  0 ],
       [ 5.4,  3.7,  0 ],
       [ 4.8,  3.4,  0 ],
       [ 4.8,  3. ,  0 ],
       [ 4.3,  3. ,  0 ],
       [ 5.8,  4. ,  0 ],
       [ 5.7,  4.4,  0 ],
       [ 5.4,  3.9,  0 ],
       [ 5.1,  3.5,  0 ],
       [ 5.7,  3.8,  0 ],
       [ 5.1,  3.8,  0 ],
       [ 7. ,  3.2,  1 ],
       [ 6.4,  3.2,  1 ],
       [ 6.9,  3.1,  1 ],
       [ 5.5,  2.3,  1 ],
       [ 6.5,  2.8,  1 ],
       [ 5.7,  2.8,  1 ],
       [ 6.3,  3.3,  1 ],
       [ 4.9,  2.4,  1 ],
       [ 6.6,  2.9,  1 ],
       [ 5.2,  2.7,  1 ],
       [ 5. ,  2. ,  1 ],
       [ 5.9,  3. ,  1 ],
       [ 6. ,  2.2,  1 ],
       [ 6.1,  2.9,  1 ],
       [ 5.6,  2.9,  1 ],
       [ 6.7,  3.1,  1 ],
       [ 5.6,  3. ,  1 ],
       [ 5.8,  2.7,  1 ],
       [ 6.2,  2.2,  1 ],
       [ 5.6,  2.5,  1 ],
       [ 5.9,  3.2,  1 ],
       [ 6.1,  2.8,  1 ],
       [ 6.3,  2.5,  1 ],
       [ 6.1,  2.8,  1 ],
       [ 6.4,  2.9,  1 ]]

data = np.array(data)
```

Plot the data:

<!-- grade_id: cell-036364c621db11d5, points: 0, grade: true, solution: true -->
```python

### visualize data ###

plt.axis([0, 10, 0, 6])
plt.grid()
for point in data:
    # your code here
    # use plt.scatter(...) function
    pass


# YOUR CODE HERE
raise NotImplementedError()
```

Set the learning rate `eta` to 1, initialize the weights `w1` and `w2` by using the `np.random.rand()` method, and the bias `b` to 0. Then, run for 100 epochs the perceptron algorithm:
- remember to shuffle the data at each epoch using the `np.random.shuffle(data)` function
- in the for loop there's already a snippet for updating the image everytime:
```
plot_border(w1, w2, b, data_norm, fig)
display.clear_output(wait=True)
time.sleep(0.05)
```
- The weights here are on the form `w1` and `w2`. However, to use your previous code and so the dot product between W and the data, you need to pack `W = np.array([w1, w2])` and then to unpack `w1, w2 = W` before plotting at each epoch!


<!-- grade_id: cell-b23508422bcf31b6, points: 0, grade: false, solution: true -->
```python
eta = 1
epochs = 100
w1 = 0
w2 = 0
b = 0
fig = plt.figure()

for epoch in range(epochs):
    # shuffle data
    for x1, x2, y in data:
        # remember to pack w1,w2 into a numpy array W
        # remember to pack x1,x2 into a numpy array W
        pass
    # plot_border(w1, w2, b, data, fig)
    # display.clear_output(wait=True)

    
    
# YOUR CODE HERE
raise NotImplementedError()
plot_border(w1, w2, b, data, fig)

```

<!-- grade_id: cell-99b2cdb1e86fbbb5, points: 1, grade: true, solution: false -->
```python
assert W.shape == (2,)
assert W[0] == w1
assert W[1] == w2
```

Implement a function `guess_flower(sepal_length, sepal_width, W, b)` that prints Iris-setosa in case of negative score and return 0, Iris-versicolor in case of positive score and return 1.

<!-- grade_id: cell-42542b2b2667947c, points: 0, grade: false, solution: true -->
```python
# prediction function: 0 for Iris-setosa and 1 for Iris-versicolor

def guess_flower(sepal_length, sepal_width, W, b):
    if 0:
        print('Iris-setosa')
    else:
        print('Iris-versicolor')
        
# YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: cell-673093aba4d0b46b, points: 1, grade: true, solution: false -->
```python
assert guess_flower(5.4, 3.4, W, b) == 0 or guess_flower(5.4, 3.4, W, b) == 1
```

```python
test_data = [[ 5.4,  3.4,  0. ],
           [ 4.6,  3.6,  0. ],
           [ 5.1,  3.3,  0. ],
           [ 4.8,  3.4,  0. ],
           [ 5. ,  3. ,  0. ],
           [ 5. ,  3.4,  0. ],
           [ 5.2,  3.5,  0. ],
           [ 5.2,  3.4,  0. ],
           [ 4.7,  3.2,  0. ],
           [ 4.8,  3.1,  0. ],
           [ 5.4,  3.4,  0. ],
           [ 5.2,  4.1,  0. ],
           [ 5.5,  4.2,  0. ],
           [ 4.9,  3.1,  0. ],
           [ 5. ,  3.2,  0. ],
           [ 5.5,  3.5,  0. ],
           [ 4.9,  3.1,  0. ],
           [ 4.4,  3. ,  0. ],
           [ 5.1,  3.4,  0. ],
           [ 5. ,  3.5,  0. ],
           [ 4.5,  2.3,  0. ],
           [ 4.4,  3.2,  0. ],
           [ 5. ,  3.5,  0. ],
           [ 5.1,  3.8,  0. ],
           [ 4.8,  3. ,  0. ],
           [ 5.1,  3.8,  0. ],
           [ 4.6,  3.2,  0. ],
           [ 5.3,  3.7,  0. ],
           [ 5. ,  3.3,  0. ],
           [ 6.8,  2.8,  1. ],
           [ 6.7,  3. ,  1. ],
           [ 6. ,  2.9,  1. ],
           [ 5.7,  2.6,  1. ],
           [ 5.5,  2.4,  1. ],
           [ 5.5,  2.4,  1. ],
           [ 5.8,  2.7,  1. ],
           [ 6. ,  2.7,  1. ],
           [ 5.4,  3. ,  1. ],
           [ 6. ,  3.4,  1. ],
           [ 6.7,  3.1,  1. ],
           [ 6.3,  2.3,  1. ],
           [ 5.6,  3. ,  1. ],
           [ 5.5,  2.5,  1. ],
           [ 5.5,  2.6,  1. ],
           [ 6.1,  3. ,  1. ],
           [ 5.8,  2.6,  1. ],
           [ 5. ,  2.3,  1. ],
           [ 5.6,  2.7,  1. ],
           [ 5.7,  3. ,  1. ],
           [ 5.7,  2.9,  1. ],
           [ 6.2,  2.9,  1. ],
           [ 5.1,  2.5,  1. ],
           [ 5.7,  2.8,  1. ]]
```

Test your model on the test set. You should be able to achieve at least 95% accuracy. Provide your model in `W` and `b` and your function `guess_flower(x1, x2, W, b)` should accept the parameters in that order.

<!-- grade_id: cell-aaa30ec08a34fef2, points: 0, grade: false, solution: true -->
```python
# If you called W and b in a different way, here you can rename it.
# Otherwise, check the accuracy you achieved

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-76baa44f7564d8d4, points: 1, grade: true, solution: false -->
```python
assert correct/total >= 0.95
```

## Challenge

Can you make a cool animation of the running algorithm and post it on the social? 😉

If you create it directly using Python, it's a plus!

```python

```